<a id="home"></a>
## Assignment 2 - KNN
This assignment has ... parts. Follow the instructions in each section.

| Section| About | Section-name | points | 
| :- | :- | :- |  :- | 
| (i) | Learning  Materials |  [Materials to review for this assignment](#review_materials)
| (ii) | Guidelines | [General assignment instructions](#gen_instructions) | 
| 1 |  Imports | [Import packages](#Imports) | 
| 2 |  Your personal details | [Student info](#student_info) | 0.5 |
| 3 |  Get Task Data | [info: YOUR personal implementation tasks](#your_implementation) | 
| 4 |  Load Data Set | [Load the dataset](#load_dataset) | 1.5 | 
| 5 |  Scale the Data | [Scaling](#scaling) | 2.5 |
| 6 |  Calculate Distances | [Calculate Test,Train Distance](#dist) | 2 | 
| 7 |  Predict | [KNN main flow - predict](#predict) | 2 |
| 8 |  Model Evaluation | [Model Evaluation](#evaluation) | 1.5 |
| - | Total Points | | 10 |

<a id="review_materials"></a>
### (i) Materials to review for this assignment
<h4>From Moodle:</h4> 
<h5><u>Review the notebooks regarding the following python topics</u>:</h5>
<div class="alert alert-info">
&#x2714; Brief review: <b>Basic Python</b>.<br/>
&#x2714; Thorough review: <b>The NumPy module</b><br/>
&#x2714; Thorough review: <b>The Pandas module</b> (both parts)<br/>
</div> 
<h5><u>Review the presentations regarding the following topics</u>:</h5>
<div class="alert alert-info">
&#x2714; <b>Statistical intro. and scaling presentation</b><br/>
&#x2714; <b>The KNN algorithm presentation</b><br/>
&#x2714; <b>Naive Bayes and model evaluation presentation</b>(you need the model evaluation part here)<br/>
</div>

[Go to the beginning of the assignment](#home)
<a id="gen_instructions"></a>
### (ii) General assignment instructions
#### Working on the assignment, todo and not todo:
<b><u>Todo</u></b>:<br />
<div class="alert alert-success">
<ul>
<li>&#x2714; Use <b> Jupyter Notebook </b>.</li>
<li>&#x2714; <b> Work by yourself</b> and submit your own assignment, <b>no pairing</b> to other students</li>
<li>&#x2714; Test and save your assignment - <b>submit the last tested and saved version</b></li>
<li>&#x2714; Work on the <b>original notebook</b> of the assignment</li>
<li>&#x2714; Run the code cells by their original order</li>
<li>&#x2714; submit <b>only</b> the <b>notebook (Assignment2_KNN.ipynb)</b> file (no name changing)</li></ul>
</div>
<b><u>Not todo</u></b>:<br />
<div class="alert alert-danger">
<ul>
<li>&#x2718; <b>do NOT use Google Collaboratory</b> </li>
<li>&#x2718; do <b>NOT</b> submit an <b>empty assignment</b></li>
<li>&#x2718; do <b>NOT submit extra files</b>, <u>unless you're asked to do so</u></li>
<li>&#x2718; <b>Do NOT submit a .py/.txt/.rar/.zip (or any non (.ipynb) file)</b> <u>version for the notebook</u> of the assignment</li>
<li>&#x2718; <b>do NOT change</b> the notebook file name</li>
</ul>
</div>

<b>Your implementations are performed in the 'Answer code cell's as following</b>:<br/>
<div class="alert alert-success">
&#x2714; Your code starts one line after the '<code># YOUR CODE HERE</code>' remark<br/>
&#x2714; Implemented code<br/>
&#x2714; returned expected value<br/><br/>
<b><u>Your code should look like this</u><b>:</div>
<code># 5. function_name (your solution)
# ====>>> YOUR ACTIONS: IMPLEMENT & RUN 

def method_to_implement():
    # YOUR CODE HERE
    your implemented code ...
</code>

[Go to the beginning of the assignment](#home)
<a id="Imports"></a>
### 1. import modules (packages)
This step is necessary in order to use external modules (packages). <br/>
<div class="alert alert-success">
&#x2714; Any module or package imported in the following code cell are permitted
</div>
<div class="alert alert-danger">
&#x2718; Any other modules or packages are ILLEGAL
</div>

##### 1. IMPORT MODULES - Run the following code cell:

```python
# 1. IMPORT MODULES 
# ====>>> YOUR ACTIONS: RUN ONLY

import sys                       # for testing use only
import os                        # for testing use only
from datetime import datetime    # for testing use only
import random                    # for testing use only
import hashlib                   # for testing use only
import pandas as pd
import numpy as np
import math 
import statistics

# --------------------------------------
# show several outputs in one cell. 
from IPython.core.interactiveshell import InteractiveShell
InteractiveShell.ast_node_interactivity = "all"
# --------------------------------------
```

[Go to the beginning of the assignment](#home)
<a id="student_info"></a>
### 2. Student information
##### methods to implement:
* my_name
* my_id

#### 2.a. Function: my_name
<a id="myName"></a>
Implement the function to return your full name (as a string), for example:<br/>
<code>    return 'John Smith'</code>
##### Implement and Run the following code cell:

<!-- grade_id: solution2a-my_name, points: 0, grade: false, solution: true -->
```python
# 2.a. my_name (your solution)
# ====>>> YOUR ACTIONS: IMPLEMENT & RUN 

def my_name():
    return 'Bar Salem'
    
```

##### 2.a.  my_name - OPTIONAL SELF TEST -  Implement and Run the following code cell (IF NEEDED):

<!-- grade_id: cell-1ff31f8fe6136755, points: 0, grade: false, solution: true -->
```python
# 2.a. my_name (optional self testing)
# ====>>> YOUR ACTIONS: (optional) IMPLEMENT & RUN 

student_name = my_name()
# --- add additional code to check your code if needed:
# YOUR CODE HERE
print(student_name)
```

Output:
```
Bar Salem
```

#### 2.b. Function: my_id
<a id="myName"></a>
Implement the function to return your ID number (as an integer number).<br/>
Note: the **id number should be equal to the one in Moodle**, for example:<br/>
<code>    return 132468579</code>
##### Implement and Run the following code cell:

<!-- grade_id: solution2b-my_id, points: 0, grade: false, solution: true -->
```python
# 2.b. my_id (your solution)
# ====>>> YOUR ACTIONS: IMPLEMENT & RUN 

def my_id():
    return 207351784
    
```

##### 2.b.  my_id - OPTIONAL SELF TEST -  Implement and Run the following code cell (IF NEEDED):

<!-- grade_id: cell-954293a143438460, points: 0, grade: false, solution: true -->
```python
# 2.b. my_id (optional self testing)
# ====>>> YOUR ACTIONS: (optional) IMPLEMENT & RUN 

student_id = my_id()
# --- add additional code to check your code if needed:
# YOUR CODE HERE
print(student_id)
```

Output:
```
207351784
```

<!-- grade_id: test2-1_student_info, points: 0.5, grade: true, solution: false -->
```python
# 2. my_name, my_id (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("2. 'my_name' and 'my_id' - Test 1 (name: test2-1_student_info, points: 0.5)")

# dataframe for output:
dt1 = datetime.now()
try:
    student_name, student_id = my_name(), my_id()
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"

s_datetime =  [dt1.strftime('%Y-%m-%d %H:%M:%S'), datetime.now().strftime('%Y-%m-%d %H:%M:%S')]
df_output = pd.DataFrame({'value': [student_name, student_id], 'type': [type(student_name), type(student_id)], 'date_time': s_datetime},index=['student_name','student_id'])

print ("\nGood Job!\nYou've passed the test for the implementation of 'my_name' and 'my_id'  :-)")

print ('\nOutput dataframe:')
print ('-----------------')
df_output
```

Output:
```
2. 'my_name' and 'my_id' - Test 1 (name: test2-1_student_info, points: 0.5)

Good Job!
You've passed the test for the implementation of 'my_name' and 'my_id'  :-)

Output dataframe:
-----------------
                  value           type            date_time
student_name  Bar Salem  <class 'str'>  2022-03-20 11:09:05
student_id    207351784  <class 'int'>  2022-03-20 11:09:05
```

[Go to the beginning of the assignment](#home)
<a id="your_implementation"></a>
### 3. Get Task Data: Your Tasks & PERSONAL IMPLEMENTATIONS
##### Run the following code cell:

```python
# 3. Get task data (Your tasks & PESRONAL IMPLEMENTATION INFO)
# ====>>> YOUR ACTIONS: RUN ONLY

def get_assignment_params(student_id):
    num_modulo = 2 ** 31-1
    created_1st_state = int(hashlib.md5(str(student_id).encode('utf-8')).hexdigest(),16) % num_modulo
    created_2nd_state = int(hashlib.md5(str(created_1st_state).encode('utf-8')).hexdigest(),16) % num_modulo
    created_3rd_state = int(hashlib.md5(str(created_2nd_state).encode('utf-8')).hexdigest(),16) % num_modulo
    created_4th_state = int(hashlib.md5(str(created_3rd_state).encode('utf-8')).hexdigest(),16) % num_modulo
    
    scale_types, dist_methods, eval_metrics = ['standardization', 'minmax[0,1]', 'minmax[-1,1]'], ['manhattan', 'euclidean', 'chebyshev'], ['accuracy', 'error_rate', 'precision', 'recall']
    
    score_cols, o_cols, binary_cols = ['math_score', 'reading_score', 'writing_score'], ['social_grp', 'parent_edu'], np.array(['gender_num', 'lunch_type', 'has_preparations'])
    
    assignment_params = {}
    
    random.seed(created_1st_state)
    y_col = random.choice(binary_cols)
    
    assignment_params['num_train1'] = 800
    assignment_params['num_train2'] = 80
    
    assignment_params['y_col'] = y_col
    assignment_params['binary_cols'] = list(binary_cols[binary_cols!=y_col])
    assignment_params['score_cols'], assignment_params['o_cols'] = score_cols, o_cols
    
    random.seed(created_2nd_state)  
    assignment_params['scale_type'] = random.choice(scale_types)
    
    random.seed(created_3rd_state)  
    assignment_params['dist_method'] = random.choice(dist_methods)
    
    random.seed(created_4th_state)  
    assignment_params['eval_metric'] = random.choice(eval_metrics)
    
    return assignment_params
# ---------------------------
try:    
    student_name, student_id = my_name(), my_id()
except Exception as e:
    print ('You probably did not implement student-info functions, \nerror Message:',str(e))
    raise
assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"         
# ---------------------------
assignment_params = get_assignment_params(student_id)
# ---------------------------
print ('Assignment 2 - KNN (- 10 points for the assignments):')
print ('-----------------------')
print ('What do you need to implement?')
print ('1.   imports - run only (to load python modules)')
print ('2.   methods: my_name, my_id - Your personal information (- 0.5 points for the test)')
print ('3.   personal implementation taks - run only')
print ('4.   methods: load_dataset - load the dataset (- 1 points for the tests)')
print ("5.   methods: scale_fit_transform, scale_transform_for_test \n\t- the scaling type YOU NEED to implement is: '" + assignment_params['scale_type'] + "' (- 2.5 points for the tests)")
print ("6. methods: calc_distance \n\t- the distance method YOU NEED to implement is: '" + assignment_params['dist_method'] + "' (- 2.5 points for the tests)")
print ("7. methods: predict - KNN predict main flow" + " (- 2 points for the tests)")
print ("8.   methods: eval_performance \n\t- the evaluation metric YOU NEED to implement is: '" + assignment_params['eval_metric'] + "' (- 1.5 points for the tests)")
```

Output:
```
Assignment 2 - KNN (- 10 points for the assignments):
-----------------------
What do you need to implement?
1.   imports - run only (to load python modules)
2.   methods: my_name, my_id - Your personal information (- 0.5 points for the test)
3.   personal implementation taks - run only
4.   methods: load_dataset - load the dataset (- 1 points for the tests)
5.   methods: scale_fit_transform, scale_transform_for_test 
	- the scaling type YOU NEED to implement is: 'minmax[-1,1]' (- 2.5 points for the tests)
6. methods: calc_distance 
	- the distance method YOU NEED to implement is: 'euclidean' (- 2.5 points for the tests)
7. methods: predict - KNN predict main flow (- 2 points for the tests)
8.   methods: eval_performance 
	- the evaluation metric YOU NEED to implement is: 'recall' (- 1.5 points for the tests)
```

[Go to the beginning of the assignment](#home)
<a id="load_dataset"></a>
### 4. Load dataset
##### methods to implement:
* load_dataset

### 4. load_dataset
<pre>
input parameters:
- file_name - a path containing the csv file name. Use the 'file_name' as is
- category_col_name
------------
The function should return a dataframe containing the feature vectors (X) and 
    a series containing the categories (y), for example:</pre>
<code>return X, y</code>

##### Implement and Run the following code cell:

<!-- grade_id: solution4-load_dataset, points: 0, grade: false, solution: true -->
```python
# 4. load_dataset (your solution)
# ====>>> YOUR ACTIONS: IMPLEMENT & RUN

def load_dataset(file_name, category_col_name):
    df=pd.read_csv(file_name)
    y=df[category_col_name]
    X=df.copy()
    X.drop(category_col_name,axis=1,inplace=True)
    return X,y
    
```

##### 4.  load_dataset - OPTIONAL SELF TEST - Implement and Run the following code cell (IF NEEDED):

<!-- grade_id: cell-1292bb4085725f5e, points: 0, grade: false, solution: true -->
```python
# 4. load_dataset (optional self testing)
# ====>>> YOUR ACTIONS: (optional) IMPLEMENT & RUN 

file_name = 'data' + os.sep + 'Students-Performance_shuffled_1000.csv'
student_name, student_id = my_name(), my_id()
assignment_params = get_assignment_params(student_id)
X, y = load_dataset(file_name, assignment_params['y_col'])
# --- add additional code to check your code if needed:
# YOUR CODE HERE
X.info()

```

Output:
```
<class 'pandas.core.frame.DataFrame'>
RangeIndex: 1000 entries, 0 to 999
Data columns (total 7 columns):
 #   Column         Non-Null Count  Dtype
---  ------         --------------  -----
 0   social_grp     1000 non-null   int64
 1   gender_num     1000 non-null   int64
 2   parent_edu     1000 non-null   int64
 3   lunch_type     1000 non-null   int64
 4   math_score     1000 non-null   int64
 5   reading_score  1000 non-null   int64
 6   writing_score  1000 non-null   int64
dtypes: int64(7)
memory usage: 54.8 KB
```

<!-- grade_id: test4-1_load_dataset, points: 0.1, grade: true, solution: false -->
```python
# 4. load_dataset (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("4. 'load_dataset' - Test 1 (name: test4-1_load_dataset, points: 0.1) - Sanity")

file_name = 'data' + os.sep + 'Students-Performance_shuffled_1000.csv'

try:    
    student_name, student_id = my_name(), my_id()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)
    X, y = load_dataset(file_name, assignment_params['y_col'])
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise
print ("Good Job!\nYou've passed the 1st test for the 'load_dataset' function implementation :-)")
```

Output:
```
4. 'load_dataset' - Test 1 (name: test4-1_load_dataset, points: 0.1) - Sanity
Good Job!
You've passed the 1st test for the 'load_dataset' function implementation :-)
```

<!-- grade_id: test4-2_load_dataset, points: 0.5, grade: true, solution: false -->
```python
# 4. load_dataset (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("4. 'load_dataset' - Test 2 (name: test4-2_load_dataset, points: 0.5)")

file_name = 'data' + os.sep + 'Students-Performance_shuffled_1000.csv'


try:    
    student_name, student_id = my_name(), my_id()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)
    X, y = load_dataset(file_name, assignment_params['y_col'])
    # ---------------------------
    X_train, y_train = X.iloc[:assignment_params['num_train1'],:],  y.iloc[:assignment_params['num_train1']]
    X_test, y_test = X.iloc[assignment_params['num_train1']:,:], y.iloc[assignment_params['num_train1']:]
    # ---------------------------
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise
assert type(X) is pd.DataFrame or type(y) is pd.Series, "Wrong type for feature vectors dataframe or for category series"

print ("Good Job!\nYou've passed the 2nd test for the 'load_dataset' function implementation :-)")
```

Output:
```
4. 'load_dataset' - Test 2 (name: test4-2_load_dataset, points: 0.5)
Good Job!
You've passed the 2nd test for the 'load_dataset' function implementation :-)
```

<!-- grade_id: test4-3_load_dataset, points: 0.9, grade: true, solution: false -->
```python
# 4. load_dataset (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("4. 'load_dataset' - Test 3 (name: test4-3_load_dataset, points: 0.9)")

file_name = 'data' + os.sep + 'Students-Performance_shuffled_1000.csv'

try:    
    student_name, student_id = my_name(), my_id()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)
    X, y = load_dataset(file_name, assignment_params['y_col'])
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

assert X.shape == (1000,7) and X.shape[0] == y.shape[0], 'Wrong shape for feature vector dataframe or category series'
assert assignment_params['y_col'] not in X.columns, "'%s' should not be part of columns, it is the category column" %(assignment_params['y_col'])
print ("Good Job!\nYou've passed the 3rd test for the 'load_dataset' function implementation :-)")
```

Output:
```
4. 'load_dataset' - Test 3 (name: test4-3_load_dataset, points: 0.9)
Good Job!
You've passed the 3rd test for the 'load_dataset' function implementation :-)
```

[Go to the beginning of the assignment](#home)
<a id="scaling"></a>
### 5. Scaling
##### methods to implement:
* scale_fit_transform
* scale_transform_for_test

##### 5. Scaling - PERSONAL IMPLEMENTATION - Run the following code cell:

```python
# 5. SCALING (personal implementation info)
# ====>>> YOUR ACTIONS: RUN ONLY

try:    
    student_name, student_id = my_name(), my_id()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"    
    assignment_params = get_assignment_params(student_id)
except Exception as e:
    print ('You probably did not implement student-info functions, \nerror Message:',str(e))
    raise
# ----------------------------

print ("The scaling type you need to implement is: '" + assignment_params['scale_type'] + "'")
```

Output:
```
The scaling type you need to implement is: 'minmax[-1,1]'
```

<b>5. Scaling (scale_fit_transform, scale_transform_for_test) - Permitted and illegal methods</b>:<br />
<div class="alert alert-success">
&#x2714; You can use any function that calculates: minimum (<code>min</code>), maximum (<code>max</code>), average (<code>mean</code>), variance (<code>var</code>) or standard deviation (<code>std</code>)<br/>
&#x2714; Any other basic mathematical function (+,*, ...)<br/>
</div>
<div class="alert alert-danger">
&#x2718; DON'T USE any function, which perform the scaling
</div>
<br>
<b><u>Important Notes - for those who need to implement (t-distribution) standardization</u></b><br/>

* <u>By default, numpy's</u> `var` & `std` <u> methods refer to the population</u>
* <u>By default, pandas'</u> `var` & `std` <u> methods refer to the sample</u>
* Use the relevant methods (& parameters if needed)

### 5.a. scale_fit_transform
<pre>
input parameters:
- X_train - the input dataframe (only feature vectors)
------------
The function should return a <b>copy</b> of the train feature vector dataframe (X_train_scaled)
    a dataframe containing scaling information (scaling_info).
---------------------
</pre>
For example, if your input `X_train` is:

| | sepal_length | sepal_width | petal_length | 
| :- | :- | :- | :- | 
| <b>0</b> |  5.1 | 3.5 | 1.4 |  
| <b>1</b> |  4.9 | 3 | 1.4 |   
| <b>2</b> |  4.7 | 3.2 | 1.3 |  
| <b>3</b> |  4.6 | 3.1 | 1.5 | 
| <b>4</b> |  5 | 3.6 | 1.4 | 

------

Cases (a) & (b): If YOU are assigned with a <b>'scale_type'</b> of <b>'minmax[0,1]'</b> or <b>'minmax[-1,1]'</b>- <br />
Then, the expected `scaling_info` output will be:

| | min | max |
| :- | :- | :- |
| sepal_length | 4.6 | 5.1 |
| sepal_width | 3.0 | 3.6 |
| petal_length | 1.3 | 1.5 |

Case (a): If YOU are assigned with a <b>'scale_type'</b> of <b>'minmax[0,1]'</b> - <br />
Then, the expected `X_train_scaled` output will be:

| | sepal_length | sepal_width | petal_length | 
| :- | :- | :- | :- |
| <b>0</b> | 1.0 | 0.833333 | 0.5 |
| <b>1</b> | 0.6 | 0.0 | 0.5 |
| <b>2</b> | 0.2 | 0.333333 | 0.0 |
| <b>3</b> | 0.0 | 0.166667 | 1.0 |
| <b>4</b> | 0.8 | 1.0 | 0.5 |

Case (b): If YOU are assigned with a <b>'scale_type'</b> of <b>'minmax[-1,1]'</b> - <br />
Then, the expected `X_train_scaled` output will be:

| | sepal_length | sepal_width | petal_length | 
| :- | :- | :- | :- |
| <b>0</b> | 1.0 | 0.66666 | 0.0 |
| <b>1</b> | 0.2 | -1.0 | 0.0 |
| <b>2</b> | -0.6 | -0.33333 | -1.0 |
| <b>3</b> | -1.0 | -0.66666 | 1.0 |
| <b>4</b> | 0.6 | 1.0 | 0.0 |

------

Case (c): If YOU are assigned with a <b>'scale_type'</b> of <b>'standardization'</b> - <br />
Then, the expected `scaling_info` output will be:

| | mean | std |
| :- | :- | :- |
| sepal_length | 4.86 | 0.207364 |
| sepal_width | 3.28 | 0.258844 |
| petal_length | 1.40 | 0.070711 |
		
And the `X_train_scaled` output will be:

| | sepal_length | sepal_width | petal_length | 
| :- | :- | :- | :- |
| <b>0</b> | 1.157383 | 0.849934 | 0.000000 |
| <b>1</b> | 0.192897 |-1.081734 | 0.000000 |
| <b>2</b> | -0.771589 | -0.309067 | -1.414214 |
| <b>3</b> | -1.253831 | -0.695401 | 1.414214 |
| <b>4</b> | 0.675140 | 1.236268 | 0.000000 |

----

And an example of the output statement:<br/>
<code>    return scaling_info, X_train_scaled</code>

##### 5.a. scale_fit_transform - OPTIONAL  ASSISTANCE  - Implement & Run the following cell (IF NEEDED):

<!-- grade_id: assistance5a-scale_fit_transform, points: 0, grade: false, solution: true -->
```python
# 5.a. scale_fit_transform (optional assitance for your solution)
# ====>>> YOUR ACTIONS: (optional) IMPLEMENT & RUN 

# ---- Add assistance code here IF NEEDED:
# YOUR CODE HERE

```

##### 5.a. scale_fit_transform - Implement and Run the following code cell:

<!-- grade_id: solution5a-scale_fit_transform, points: 0, grade: false, solution: true -->
```python
# 5.a. scale_fit_transform (your solution)
# ====>>> YOUR ACTIONS: IMPLEMENT & RUN
def scale_fit_transform(X_train):
    X_train_scaled=X_train.copy()
    X_train_copy = X_train_scaled.agg(['min','max'])
    scaling_info=X_train_copy.T
    i=0
    for col in X_train:
        minValue=scaling_info.iloc[i,0]
        colRange=scaling_info.iloc[i,1]-minValue
        i+=1
        X_train_scaled[col]-=minValue
        X_train_scaled[col]=(X_train_scaled[col]/colRange)*2-1
    return scaling_info, X_train_scaled
    
```

##### 5.a. scale_fit_transform - OPTIONAL SELF TEST - Implement & Run the following code cell (IF NEEDED):

<!-- grade_id: cell-b95f59dfb75314fd, points: 0, grade: false, solution: true -->
```python
# 5.a. scale_fit_transform (optional self testing)
# ====>>> YOUR ACTIONS: (optional) IMPLEMENT & RUN 

file_name = 'data' + os.sep + 'Students-Performance_shuffled_1000.csv'
y_col     = 'has_preparations'
student_name, student_id = my_name(), my_id()
assignment_params = get_assignment_params(student_id)    
X, y = load_dataset(file_name, y_col)
X_train, y_train = X.iloc[:assignment_params['num_train1'],:],  y.iloc[:assignment_params['num_train1']]
X_test, y_test = X.iloc[assignment_params['num_train1']:,:], y.iloc[assignment_params['num_train1']:]
trained_scaling_info,  X_train_scaled = scale_fit_transform(X_train)
# --- add additional code to check your code if needed:
trained_scaling_info

```

Output:
```
min  max
social_grp       1    5
gender_num       0    1
parent_edu       0    5
lunch_type       0    1
math_score      16  100
reading_score   23  100
writing_score   24  100
```

##### 5.a. scale_fit_transform - GRADED TESTS -  Run the following (3) code cells:

<!-- grade_id: test5a-1_scale_fit_transform, points: 0.1, grade: true, solution: false -->
```python
# 5.a. scale_fit_transform (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("5.a. 'scale_fit_transform' - Test 1 (name: test5a-1_scale_fit_transform, points: 0.1) - Sanity")

file_name = 'data' + os.sep + 'Students-Performance_shuffled_1000.csv'
y_col     = 'has_preparations'
try:    
    student_name, student_id = my_name(), my_id()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)    
    print ("\t\tTesing a '" + assignment_params['scale_type'] + "' scale type ...\n")
    X, y = load_dataset(file_name, y_col)
    X_train, y_train = X.iloc[:assignment_params['num_train1'],:],  y.iloc[:assignment_params['num_train1']]
    X_test, y_test = X.iloc[assignment_params['num_train1']:,:], y.iloc[assignment_params['num_train1']:]
    trained_scaling_info,  X_train_scaled = scale_fit_transform(X_train)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise


print ("\nGood Job!\nYou've passed the 1st test for the 'scale_fit_transform' function implementation :-)")
```

Output:
```
5.a. 'scale_fit_transform' - Test 1 (name: test5a-1_scale_fit_transform, points: 0.1) - Sanity
		Tesing a 'minmax[-1,1]' scale type ...


Good Job!
You've passed the 1st test for the 'scale_fit_transform' function implementation :-)
```

<!-- grade_id: test5a-2_scale_fit_transform, points: 0.5, grade: true, solution: false -->
```python
# 5.a. scale_fit_transform (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("5.a. 'scale_fit_transform' - Test 2 (name: test5a-2_scale_fit_transform, points: 0.5)")

file_name = 'data' + os.sep + 'Students-Performance_shuffled_1000.csv'
y_col     = 'has_preparations'
try:      
    student_name, student_id = my_name(), my_id()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)    
    print ("\t\tTesing a '" + assignment_params['scale_type'] + "' scale type ...\n")
    X, y = load_dataset(file_name, y_col)
    X_train, y_train = X.iloc[:assignment_params['num_train1'],:],  y.iloc[:assignment_params['num_train1']]
    X_test, y_test = X.iloc[assignment_params['num_train1']:,:], y.iloc[assignment_params['num_train1']:]
    trained_scaling_info,  X_train_scaled = scale_fit_transform(X_train)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

check_equal = lambda itr_a,itr_b: len(list(itr_a))>0 and len(list(itr_a))==len(list(itr_b)) and 0==len(set(list(itr_a))-set(list(itr_b))) and 0==len(set(list(itr_b))-set(list(itr_a))) and len(list(itr_a))==len(set(list(itr_b))|set(list(itr_b)))
attrs=assignment_params['o_cols']+assignment_params['score_cols']+['gender_num', 'lunch_type']

assert X_train_scaled.shape == X_train.shape, 'Wrong shape for scaled train dataframe'
np.testing.assert_array_equal(X_train_scaled.index, X_train.index, 'scaled train and train dataframes should have the same index')
assert check_equal(X_train_scaled.columns, attrs), 'scaled trained and train dataframes should have the same column names'

print ("\nGood Job!\nYou've passed the 2nd test for the 'scale_fit_transform' function implementation :-)")
```

Output:
```
5.a. 'scale_fit_transform' - Test 2 (name: test5a-2_scale_fit_transform, points: 0.5)
		Tesing a 'minmax[-1,1]' scale type ...


Good Job!
You've passed the 2nd test for the 'scale_fit_transform' function implementation :-)
```

<!-- grade_id: test5a-3_scale_fit_transform, points: 0.8, grade: true, solution: false -->
```python
# 5.a. scale_fit_transform (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("Parts 5.a. 'scale_fit_transform' - Test 3 (name: test5a-3_scale_fit_transform, points: 0.8)")


file_name = 'data' + os.sep + 'Students-Performance_shuffled_1000.csv'
y_col     = 'has_preparations'
try:      
    student_name, student_id = my_name(), my_id()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)    
    print ("\t\tTesing a '" + assignment_params['scale_type'] + "' scale type ...\n")
    X, y = load_dataset(file_name, y_col)
    X_train, y_train = X.iloc[:assignment_params['num_train1'],:],  y.iloc[:assignment_params['num_train1']]
    X_test, y_test = X.iloc[assignment_params['num_train1']:,:], y.iloc[assignment_params['num_train1']:]
    trained_scaling_info,  X_train_scaled = scale_fit_transform(X_train)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise
mean_values = X_train_scaled.mean().values
mean_round  = np.round(mean_values+0.00001,4)
min_values = X_train_scaled.min().values
min_round  = np.round(min_values+0.00001,4)
max_values = X_train_scaled.max().values
max_round  = np.round(max_values+0.00001,4)
if 'standardization' == assignment_params['scale_type']:
    np.testing.assert_array_equal([0.,0.,0.,0.,0.,0.,0.], mean_round, 'Wrong average standardization scaled values')
elif assignment_params['scale_type'] == 'minmax[0,1]':
    np.testing.assert_array_equal([0.,0.,0.,0.,0.,0.,0.], min_round, 'Wrong min minmax[0,1] scaled values')
    np.testing.assert_array_equal([1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0], max_round, 'Wrong max minmax[0,1] scaled values')
else:    
    np.testing.assert_array_equal([-1.0, -1.0, -1.0, -1.0, -1.0, -1.0, -1.0], min_round, 'Wrong min minmax[-1,1] scaled values') 
    np.testing.assert_array_equal([1.0, 1.0, 1.0, 1.0, 1.0, 1.0, 1.0], max_round, 'Wrong max minmax[-1,1] scaled values')
print ("\nGood Job!\nYou've passed the 3rd test for the 'scale_fit_transform' functions implementation :-)")
```

Output:
```
Parts 5.a. 'scale_fit_transform' - Test 3 (name: test5a-3_scale_fit_transform, points: 0.8)
		Tesing a 'minmax[-1,1]' scale type ...


Good Job!
You've passed the 3rd test for the 'scale_fit_transform' functions implementation :-)
```

### 5.b. scale_transform_for_test
<pre>
input parameters:
- X_test - the input dataframe (only feature vectors)
- trained_scaling_info - scaling information (returned from the above scale_fit_transform method) 
------------
The function should return a scaled <b>copy</b> of the test dataframe (X_test_scaled), for example:</pre>
<code>    return X_test_scaled</code>

##### 5.b. scale_transform_for_test - OPTIONAL  ASSISTANCE  - Implement & Run the following (IF NEEDED):

<!-- grade_id: assistance5b-scale_transform_for_test, points: 0, grade: false, solution: true -->
```python
# 5.b. scale_transform_for_test (optional assitance for your solution)
# ====>>> YOUR ACTIONS: (optional) IMPLEMENT & RUN 

# ---- Add assistance code here IF NEEDED:
# YOUR CODE HERE

```

##### 5.b. scale_transform_for_test - Implement and Run the following code cell:

<!-- grade_id: solution5b-scale_transform_for_test, points: 0, grade: false, solution: true -->
```python
# 5.b. scale_transform_for_test (your solution)
# ====>>> YOUR ACTIONS: IMPLEMENT & RUN

def scale_transform_for_test(trained_scaling_info,  X_test):
    X_test_scaled=X_test.copy()
    i=0
    for col in X_test_scaled:
        minValue=trained_scaling_info.iloc[i,0]
        colRange=trained_scaling_info.iloc[i,1]-minValue
        i+=1
        X_test_scaled[col]-=minValue
        X_test_scaled[col]=(X_test_scaled[col]/colRange)*2-1
    return X_test_scaled
    
```

##### 5.b. scale_transform_for_test - OPTIONAL SELF TEST - Implement & Run the following cell (IF NEEDED):

<!-- grade_id: cell-14bfe00fb1108c30, points: 0, grade: false, solution: true -->
```python
# 5.b. scale_transform_for_test (optional self testing)
# ====>>> YOUR ACTIONS: (optional) IMPLEMENT & RUN 

file_name = 'data' + os.sep + 'Students-Performance_shuffled_1000.csv'
y_col     = 'has_preparations'
student_name, student_id = my_name(), my_id()
assignment_params = get_assignment_params(student_id)    
X, y = load_dataset(file_name, y_col)
X_train, y_train = X.iloc[:assignment_params['num_train1'],:],  y.iloc[:assignment_params['num_train1']]
X_test, y_test = X.iloc[assignment_params['num_train1']:,:], y.iloc[assignment_params['num_train1']:]
trained_scaling_info,  X_train_scaled = scale_fit_transform(X_train)
X_test_scaled = scale_transform_for_test(trained_scaling_info,  X_test)
# --- add additional code to check your code if needed:


```

##### 5.b. scale_fit_transform, scale_transform_for_test - GRADED TESTS -  Run the following (3) code cells:

<!-- grade_id: test5b-1_scale_transform_for_test, points: 0.1, grade: true, solution: false -->
```python
# 5.b. scale_transform_for_test (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("Parts 5.b. 'scale_transform_for_test' - Test 1 (name: test5b-1_scale_transform_for_test, points: 0.1) - Sanity")

file_name = 'data' + os.sep + 'Students-Performance_shuffled_1000.csv'
y_col     = 'has_preparations'
try:      
    student_name, student_id = my_name(), my_id()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)    
    print ("\t\tTesing a '" + assignment_params['scale_type'] + "' scale type ...\n")
    X, y = load_dataset(file_name, y_col)
    X_train, y_train = X.iloc[:assignment_params['num_train1'],:],  y.iloc[:assignment_params['num_train1']]
    X_test, y_test = X.iloc[assignment_params['num_train1']:,:], y.iloc[assignment_params['num_train1']:]
    trained_scaling_info,  X_train_scaled = scale_fit_transform(X_train)
    X_test_scaled = scale_transform_for_test(trained_scaling_info,  X_test)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise
print ("\nGood Job!\nYou've passed the 1st test for the 'scale_transform_for_test' function implementation :-)")
```

Output:
```
Parts 5.b. 'scale_transform_for_test' - Test 1 (name: test5b-1_scale_transform_for_test, points: 0.1) - Sanity
		Tesing a 'minmax[-1,1]' scale type ...


Good Job!
You've passed the 1st test for the 'scale_transform_for_test' function implementation :-)
```

<!-- grade_id: test5b-2_scale_transform_for_test, points: 0.5, grade: true, solution: false -->
```python
# 5.b. scale_transform_for_test (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("Parts 5.b. 'scale_transform_for_test' - Test 2 (name: test5b-2_scale_transform_for_test, points: 0.5)")

file_name = 'data' + os.sep + 'Students-Performance_shuffled_1000.csv'
y_col     = 'has_preparations'
try:      
    student_name, student_id = my_name(), my_id()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)    
    print ("\t\tTesing a '" + assignment_params['scale_type'] + "' scale type ...\n")
    X, y = load_dataset(file_name, y_col)
    X_train, y_train = X.iloc[:assignment_params['num_train1'],:],  y.iloc[:assignment_params['num_train1']]
    X_test, y_test = X.iloc[assignment_params['num_train1']:,:], y.iloc[assignment_params['num_train1']:]
    trained_scaling_info,  X_train_scaled = scale_fit_transform(X_train)
    X_test_scaled = scale_transform_for_test(trained_scaling_info,  X_test)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise
check_equal = lambda itr_a,itr_b: len(list(itr_a))>0 and len(list(itr_a))==len(list(itr_b)) and 0==len(set(list(itr_a))-set(list(itr_b))) and 0==len(set(list(itr_b))-set(list(itr_a))) and len(list(itr_a))==len(set(list(itr_b))|set(list(itr_b)))
assert X_test_scaled.shape == X_test.shape, 'Wrong shape for scaled test dataframe'
np.testing.assert_array_equal(X_test_scaled.index, X_test.index, 'scaled test and test dataframes should have the same index')
np.testing.assert_array_equal(X_train_scaled.columns, X_test_scaled.columns, 'scaled train and scaled test dataframes should have the same column names')

print ("\nGood Job!\nYou've passed the 2nd test for the 'scale_transform_for_test' function implementation :-)")
```

Output:
```
Parts 5.b. 'scale_transform_for_test' - Test 2 (name: test5b-2_scale_transform_for_test, points: 0.5)
		Tesing a 'minmax[-1,1]' scale type ...


Good Job!
You've passed the 2nd test for the 'scale_transform_for_test' function implementation :-)
```

<!-- grade_id: test5-scaling-train-test, points: 1, grade: true, solution: false -->
```python
# 5. scale_fit_transform, scale_transform_for_test (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("Part 5. 'scale_transform_for_test' (& 'scale_fit_transform') - Test (name: test5-scaling-train-test, points: 1.0)")

file_name = 'data' + os.sep + 'Students-Performance_shuffled_1000.csv'
y_col     = 'has_preparations'
try:      
    student_name, student_id = my_name(), my_id()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)    
    print ("\t\tTesing a '" + assignment_params['scale_type'] + "' scale type ...\n")
    X, y = load_dataset(file_name, y_col)
    X_train, y_train = X.iloc[:assignment_params['num_train1'],:],  y.iloc[:assignment_params['num_train1']]
    X_test, y_test = X.iloc[assignment_params['num_train1']:,:], y.iloc[assignment_params['num_train1']:]
    trained_scaling_info,  X_train_scaled = scale_fit_transform(X_train)
    X_test_scaled = scale_transform_for_test(trained_scaling_info,  X_test)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise
check_equal = lambda itr_a,itr_b: len(list(itr_a))>0 and len(list(itr_a))==len(list(itr_b)) and 0==len(set(list(itr_a))-set(list(itr_b))) and 0==len(set(list(itr_b))-set(list(itr_a))) and len(list(itr_a))==len(set(list(itr_b))|set(list(itr_b)))
mean_mean_scl_test = X_test_scaled.mean().values.mean()
mean_mean_scl_test_rnd =np.round(mean_mean_scl_test,4)
abs_mean_mean_scl_test = abs(mean_mean_scl_test_rnd)
max_mean_scl_test =X_test_scaled.mean().values.max()
max_mean_scl_test_rnd =np.round(max_mean_scl_test,4)
trained_scaling_info_means = trained_scaling_info.mean().values 

if 'standardization' == assignment_params['scale_type']:
    assert 0<mean_mean_scl_test_rnd and mean_mean_scl_test_rnd<0.1, 'Wrong average mean standardized scaled values'
    assert 0.04<max_mean_scl_test_rnd and mean_mean_scl_test_rnd<0.14, 'Wrong max mean standardized scaled values'
elif assignment_params['scale_type'] == 'minmax[0,1]':
    assert 0.5<mean_mean_scl_test_rnd and mean_mean_scl_test_rnd<0.6, 'Wrong average mean minmax[0,1] scaled values'
    assert 0.6<max_mean_scl_test_rnd and mean_mean_scl_test_rnd<0.7, 'Wrong max mean minmax[0,1] scaled values'
else: # assignment_params['scale_type'] == 'minmax[-1,1]':
    assert 0.07<mean_mean_scl_test_rnd and mean_mean_scl_test_rnd<0.17, 'Wrong average mean minmax[-1,1] scaled values'
    assert 0.23<max_mean_scl_test_rnd and mean_mean_scl_test_rnd<0.33, 'Wrong max mean minmax[-1,1] scaled values'
print ("\nGood Job!\nYou've passed the test for both the 'scale_fit_transform' and 'scale_transform_for_test' functions implementation :-)")
```

Output:
```
Part 5. 'scale_transform_for_test' (& 'scale_fit_transform') - Test (name: test5-scaling-train-test, points: 1.0)
		Tesing a 'minmax[-1,1]' scale type ...


Good Job!
You've passed the test for both the 'scale_fit_transform' and 'scale_transform_for_test' functions implementation :-)
```

[Go to the beginning of the assignment](#home)
<a id="dist"></a>
### 6. Calculate (vector) distance 

##### 6. calc_distance - PERSONAL IMPLEMENTATION - Run the following code cell:

```python
# 6. calc_distance (personal implementation info)
# ====>>> YOUR ACTIONS: RUN ONLY

try:    
    student_name, student_id = my_name(), my_id()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)
except Exception as e:
    print ('You probably did not implement student-info functions, \nerror Message:',str(e))
    raise

print ("The distance method you need to implement is: '" + assignment_params['dist_method'] + "'")
```

Output:
```
The distance method you need to implement is: 'euclidean'
```

<b>6. calc_distance (Calculate test-train vector distance) - Permitted and illegal methods</b>:<br />
<div class="alert alert-success">
&#x2714; You can use any function that calculates: absolute value (<code>abs</code>), power (such as <code>pow</code>), square root (such as <code>sqrt</code>) or sum (<code>sum</code>)<br/>
&#x2714; You can use any function that calculates basic vector or scalar mathematical function<br/>
</div>
<div class="alert alert-danger">
&#x2718; DON'T USE any function which performs any of the distance functions<br/>
&#x2718; DON'T USE any function which performs a closer step towards the distance functions
</div>

### 6. calc_distance
<pre>
input parameters:
- X_test  - test set feature vector dataframe
- X_train - train set feature vector dataframe
------------
The function should return a dataframe, containing the vector distances, between each of the test vectors and each of the train vectors.
  * Assuming df_dist is the output dataframe, df_dist.columns (column index) should 
      be equal to the X_train.index and df_dist.index should be equal to the X_test.index.
  * Each cell should contain the distance between the test vector, represented by the index, 
           and train vector, represented by the column name.
  * Use the distance method which is ASSIGNED TO YOU
-----
An example of a return statement:</pre>
<code>    return df_dist</code>

##### 6. calc_distance - OPTIONAL  ASSISTANCE  - Implement & Run the following code cell (IF NEEDED):

<!-- grade_id: assistance6-calc_distance, points: 0, grade: false, solution: true -->
```python
# 6. calc_distance (optional assitance for your solution)
# ====>>> YOUR ACTIONS: (optional) IMPLEMENT & RUN 

# ---- Add assistance code here IF NEEDED:
# YOUR CODE HERE

```

##### 6. calc_distance - Implement and Run the following code cell:

<!-- grade_id: solution6-calc_distance, points: 0, grade: false, solution: true -->
```python
# 6. calc_distance (your solution)
# ====>>> YOUR ACTIONS: IMPLEMENT & RUN 

def calc_distance(X_test,X_train):
    a=len(list(X_train.index))
    b=len(list(X_test.index))
    df_dist = pd.DataFrame(index=X_test.index,columns=X_train.index)
    for i in range(0,a):
        for j in range(0,b):
            df_dist.iloc[j,i]=math.sqrt(pow((X_train.iloc[i,0]-X_test.iloc[j,0]),2)+pow((X_train.iloc[i,1]-X_test.iloc[j,1]),2)+pow((X_train.iloc[i,2]-X_test.iloc[j,2]),2))
    return df_dist
    
```

##### 6. calc_distance - OPTIONAL SELF TEST - Implement & Run the following cell (IF NEEDED):

<!-- grade_id: cell-dc7df3a6321e37c9, points: 0, grade: false, solution: true -->
```python
# 6. calc_distance (optional self testing)
# ====>>> YOUR ACTIONS: (optional) IMPLEMENT & RUN 

file_name = 'data' + os.sep + 'Students-Performance_scaled_100.csv'
y_col     = 'category'

student_name, student_id = my_name(), my_id()
assignment_params = get_assignment_params(student_id)
X, y = load_dataset(file_name, y_col)
X_train, y_train = X.iloc[:assignment_params['num_train2'],:],  y.iloc[:assignment_params['num_train2']]
X_test, y_test = X.iloc[assignment_params['num_train2']:,:], y.iloc[assignment_params['num_train2']:]
dist_dataframe = calc_distance(X_test,X_train)
# --- add additional code to check your code if needed:
dist_dataframe

```

Output:
```
0         1         2         3         4         5         6   \
80  0.289321  0.222091  0.200563  0.400152  0.807372  0.585328  0.340519   
81  0.725869  0.628981  0.565881  0.768361  1.025768  0.477315  0.695094   
82  0.531697  0.594929  0.360149  0.470286  0.570787  0.308729  0.332026   
83  0.478264  0.386251  0.433914  0.577701   1.13466   0.77006   0.65914   
84  0.301623  0.230471  0.272043  0.436083  0.650466  0.620448  0.207532   
85  0.654129  0.439077  0.646223  0.825875  1.116999  0.899052  0.726543   
86  0.386627  0.347135  0.238587  0.444418  0.723075  0.456829  0.312791   
87  0.405339  0.531879  0.258392  0.304067  0.603512  0.446897  0.268992   
88  0.664338  0.456864  0.645227  0.833783  0.876559  0.818178  0.597248   
89  0.237384  0.364308  0.345779  0.329318  0.600635  0.792038  0.197582   
90  0.591855   0.52666  0.418293  0.620836  0.896439  0.386266  0.546209   
91   0.59172   0.37112  0.583773  0.767981  0.999812  0.837745  0.625779   
92  0.438046  0.345364
```

##### 6. calc_distance - GRADED TESTS -  Run the following (3) code cells:

<!-- grade_id: test6-1_calc_distance, points: 0.1, grade: true, solution: false -->
```python
# 6. calc_distance (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("6. 'calc_distance' - Test 1 (name: test6-1_calc_distance, points: 0.1) - Sanity")

file_name = 'data' + os.sep + 'Students-Performance_scaled_100.csv'
y_col     = 'category'

try:    
    student_name, student_id = my_name(), my_id()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)
    print ("\t\tTesing a '" + assignment_params['dist_method'] + "' distance method ...\n")
    X, y = load_dataset(file_name, y_col)
    X_train, y_train = X.iloc[:assignment_params['num_train2'],:],  y.iloc[:assignment_params['num_train2']]
    X_test, y_test = X.iloc[assignment_params['num_train2']:,:], y.iloc[assignment_params['num_train2']:]
    dist_dataframe = calc_distance(X_test,X_train)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise


print ("Good Job!\nYou've passed the 1st test for the 'calc_distance' function implementation :-)")
```

Output:
```
6. 'calc_distance' - Test 1 (name: test6-1_calc_distance, points: 0.1) - Sanity
		Tesing a 'euclidean' distance method ...

Good Job!
You've passed the 1st test for the 'calc_distance' function implementation :-)
```

<!-- grade_id: test6-2_calc_distance, points: 0.9, grade: true, solution: false -->
```python
# 6. calc_distance (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("6. 'calc_distance' - Test 2 (name: test6-2_calc_distance, points: 0.9)")


file_name = 'data' + os.sep + 'Students-Performance_scaled_100.csv'
y_col     = 'category'

try:    
    student_name, student_id = my_name(), my_id()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)
    print ("\t\tTesing a '" + assignment_params['dist_method'] + "' distance method ...\n")
    X, y = load_dataset(file_name, y_col)
    X_train, y_train = X.iloc[:assignment_params['num_train2'],:],  y.iloc[:assignment_params['num_train2']]
    X_test, y_test = X.iloc[assignment_params['num_train2']:,:], y.iloc[assignment_params['num_train2']:]
    dist_dataframe = calc_distance(X_test,X_train)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

assert dist_dataframe.shape[0] == X_test.shape[0] and dist_dataframe.shape[1] == X_train.shape[0], 'Mismatch shape of distance dataframes with test and train datasets'
np.testing.assert_array_equal(dist_dataframe.index, X_test.index, 'Index of distance dataframe should match test index')
np.testing.assert_array_equal(dist_dataframe.columns, X_train.index, 'Column index of distance dataframe should match training index')

print ("Good Job!\nYou've passed the 2nd test for the 'calc_distance' function implementation :-)")
```

Output:
```
6. 'calc_distance' - Test 2 (name: test6-2_calc_distance, points: 0.9)
		Tesing a 'euclidean' distance method ...

Good Job!
You've passed the 2nd test for the 'calc_distance' function implementation :-)
```

<!-- grade_id: test6-3_calc_distance, points: 1, grade: true, solution: false -->
```python
# 6. calc_distance (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("6. 'calc_distance' - Test 3 (name: test6-3_calc_distance, points: 1)")


file_name = 'data' + os.sep + 'Students-Performance_scaled_100.csv'
y_col     = 'category'

try:    
    student_name, student_id = my_name(), my_id()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)
    print ("\t\tTesing a '" + assignment_params['dist_method'] + "' distance method ...\n")
    X, y = load_dataset(file_name, y_col)
    X_train, y_train = X.iloc[:assignment_params['num_train2'],:],  y.iloc[:assignment_params['num_train2']]
    X_test, y_test = X.iloc[assignment_params['num_train2']:,:], y.iloc[assignment_params['num_train2']:]
    dist_dataframe = calc_distance(X_test,X_train)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise
mean_min_dist = dist_dataframe.min(axis=1).mean()
mean_max_dist = dist_dataframe.max(axis=1).mean()
mean_mean_dist = dist_dataframe.mean(axis=1).mean()
if assignment_params['dist_method'] == 'chebyshev':
    assert 0.7 < mean_max_dist and mean_max_dist < 0.8, "Wrong distance values for 'chebyshev' distance"
    assert 0.34 < mean_mean_dist and mean_mean_dist < 0.44, "Wrong distance values for 'chebyshev' distance"
elif assignment_params['dist_method'] == 'euclidean':
    assert 0.87 < mean_max_dist and mean_max_dist < 0.97, "Wrong distance values for 'euclidean' distance"
    assert 0.42 < mean_mean_dist and mean_mean_dist < 0.52, "Wrong distance values for 'euclidean' distance"
else: # assignment_params['dist_method'] == 'manhattan':
    assert 1.42 < mean_max_dist and mean_max_dist < 1.52, "Wrong distance values for 'manhattan' distance"
    assert 0.67 < mean_mean_dist and mean_mean_dist < 0.77, "Wrong distance values for 'manhattan' distance"
print ("Good Job!\nYou've passed the 3rd test for the 'calc_distance' function implementation :-)")
```

Output:
```
6. 'calc_distance' - Test 3 (name: test6-3_calc_distance, points: 1)
		Tesing a 'euclidean' distance method ...

Good Job!
You've passed the 3rd test for the 'calc_distance' function implementation :-)
```

[Go to the beginning of the assignment](#home)
<a id="predict"></a>
### 7. KNN predict - main flow

<b>7. predict (KNN - main flow) - Permitted and illegal methods</b>:<br />
<div class="alert alert-success">
&#x2714; You are expected to use the above 'calc_distance' function<br/>
&#x2714; You can use any function that calculates minimum, maximum, top scores, sorting and so on<br/>
&#x2714; You can use any function that calculates argmin, argmax, argsort and so on<br/>
</div>
<div class="alert alert-danger">
&#x2718; DON'T USE any function which performs KNN's main flow<br/>
&#x2718; DON'T USE any function which performs a closer step towards KNN's main flow
</div>

### 7. predict
<pre>
input parameters:
- X_test - the test set feature vector dataframe
- X_train - the train set feature vector dataframe
- y_train - the train set category series 
- k - the number of Nearest Neighbors (we select the majority out of k votes)
------------
The function should return a series consisting of the predicted category series. 
   * Use the above 'calc_distance' function, to get the distance, between 
     the test feature vector, to each of the training feature vectors.
   * the returned series, should contain an index, which is equal to the X_test index.
-----
An example of a return statement:</pre>
<code>    return y_predicted</code>

##### 7. predict - OPTIONAL  ASSISTANCE  - Implement & Run the following code cell (IF NEEDED):

<!-- grade_id: assistance7-predict, points: 0, grade: false, solution: true -->
```python
# 7. predict (optional assitance for your solution)
# ====>>> YOUR ACTIONS: (optional) IMPLEMENT & RUN 

# ---- Add assistance code here IF NEEDED:
# YOUR CODE HERE

```

##### 7. predict - Implement and Run the following code cell:

<!-- grade_id: solution7-predict, points: 0, grade: false, solution: true -->
```python
# 7. predict (your solution)
# ====>>> YOUR ACTIONS: IMPLEMENT & RUN 

def predict(X_test, X_train, y_train, k):
    distDF=calc_distance(X_test,X_train)
    cols=len(list(X_test.index))
    fTrue=0
    fFalse=0
    y_predict=[]
    for j in range(0,cols):
        smallestKDist=np.argpartition(distDF.iloc[j],k)
        for i in range(0,k):
            if(y_train[smallestKDist[i]]==1):
                fTrue+=1
            else:
                fFalse+=1
        if(fTrue>fFalse):
            y_predict.append(1)
        else:
            y_predict.append(0)
        fTrue=fFalse=0
    y_predicted=pd.Series(y_predict, index=X_test.index)
    return y_predicted
```

##### 7. predict - OPTIONAL SELF TEST - Implement & Run the following cell (IF NEEDED):

<!-- grade_id: cell-bff25b1aa05713fb, points: 0, grade: false, solution: true -->
```python
# 7. predict (optional self testing)
# ====>>> YOUR ACTIONS: (optional) IMPLEMENT & RUN 

file_name = 'data' + os.sep + 'Students-Performance_scaled_100.csv'
y_col     = 'category'

student_name, student_id = my_name(), my_id()
assignment_params = get_assignment_params(student_id)
X, y = load_dataset(file_name, y_col)
X_train, y_train = X.iloc[:assignment_params['num_train2'],:],  y.iloc[:assignment_params['num_train2']]
X_test, y_test = X.iloc[assignment_params['num_train2']:,:], y.iloc[assignment_params['num_train2']:]
y_predicted = predict(X_test, X_train, y_train, 1)
# --- add additional code to check your code if needed:
# YOUR CODE HERE
y_predicted
```

Output:
```
80    0
81    0
82    0
83    1
84    0
85    1
86    1
87    0
88    0
89    1
90    0
91    1
92    0
93    0
94    0
95    0
96    0
97    0
98    0
99    0
dtype: int64
```

##### 7. predict - GRADED TESTS -  Run the following (3) code cells:

<!-- grade_id: test7-1_predict, points: 0.1, grade: true, solution: false -->
```python
# 7. predict (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("7. 'predict' - Test 1 (name: test7-1_predict, points: 0.1) - Sanity")
file_name = 'data' + os.sep + 'Students-Performance_scaled_100.csv'
y_col     = 'category'
try:      
    student_name, student_id = my_name(), my_id()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)    
    print ("\t\tTesing a '" + assignment_params['dist_method'] + "' distance method ...\n")
    X, y = load_dataset(file_name, y_col)
    X_train, y_train = X.iloc[:assignment_params['num_train2'],:],  y.iloc[:assignment_params['num_train2']]
    X_test, y_test = X.iloc[assignment_params['num_train2']:,:], y.iloc[assignment_params['num_train2']:]
    y_predicted = predict(X_test, X_train, y_train, 1)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise
print ("Good Job!\nYou've passed the 1st test for the 'predict' function implementation :-)")     
```

Output:
```
7. 'predict' - Test 1 (name: test7-1_predict, points: 0.1) - Sanity
		Tesing a 'euclidean' distance method ...

Good Job!
You've passed the 1st test for the 'predict' function implementation :-)
```

<!-- grade_id: test7-2_predict, points: 0.9, grade: true, solution: false -->
```python
# 7. predict (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("7. 'predict' - Test 2 (name: test7-2_predict, points: 0.9)")

file_name = 'data' + os.sep + 'Students-Performance_scaled_100.csv'
y_col     = 'category'
try:      
    student_name, student_id = my_name(), my_id()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)    
    print ("\t\tTesing a '" + assignment_params['dist_method'] + "' distance method ...\n")
    X, y = load_dataset(file_name, y_col)
    X_train, y_train = X.iloc[:assignment_params['num_train2'],:],  y.iloc[:assignment_params['num_train2']]
    X_test, y_test = X.iloc[assignment_params['num_train2']:,:], y.iloc[assignment_params['num_train2']:]
    y_predicted = predict(X_test, X_train, y_train, 1)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise
assert y_test.shape==y_predicted.shape,'Wrong shape for predicted series'
np.testing.assert_array_equal(y_predicted.index, y_test.index, 'Index of predicted series should match test series')

np.testing.assert_equal(np.any(np.not_equal(y_predicted.values,y_test.values)), True,'predicted and test series should not be identicle')
print ("Good Job!\nYou've passed the 2nd test for the 'predict' function implementation :-)")
```

Output:
```
7. 'predict' - Test 2 (name: test7-2_predict, points: 0.9)
		Tesing a 'euclidean' distance method ...

Good Job!
You've passed the 2nd test for the 'predict' function implementation :-)
```

<!-- grade_id: test7-3_predict, points: 1, grade: true, solution: false -->
```python
# 7. predict (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("7. 'predict' - Test 3 (name: test7-3_predict, points: 1)")

file_name = 'data' + os.sep + 'Students-Performance_scaled_100.csv'
y_col     = 'category'
try:      
    student_name, student_id = my_name(), my_id()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)    
    print ("\t\tTesing a '" + assignment_params['dist_method'] + "' distance method ...\n")
    X, y = load_dataset(file_name, y_col)
    X_train, y_train = X.iloc[:assignment_params['num_train2'],:],  y.iloc[:assignment_params['num_train2']]
    X_test, y_test = X.iloc[assignment_params['num_train2']:,:], y.iloc[assignment_params['num_train2']:]
    y_predicted = predict(X_test, X_train, y_train, 1)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise
eq_elements = sum(1 - np.abs(y_predicted.values - y_test.values))
if assignment_params['dist_method'] == 'chebyshev':
    assert 12 < eq_elements and eq_elements < 15, "Wrong predict values for 'chebyshev' distance"
elif assignment_params['dist_method'] == 'euclidean':
    assert 14 < eq_elements and eq_elements < 17, "Wrong predict values for 'euclidean' distance"
else: # assignment_params['dist_method'] == 'manhattan':
    assert 13 < eq_elements and eq_elements < 16, "Wrong predict values for 'manhattan' distance"

print ("Good Job!\nYou've passed the 3rd test for the 'predict' function implementation :-)")
```

Output:
```
7. 'predict' - Test 3 (name: test7-3_predict, points: 1)
		Tesing a 'euclidean' distance method ...

Good Job!
You've passed the 3rd test for the 'predict' function implementation :-)
```

[Go to the beginning of the assignment](#home)
<a id="evaluation"></a>
### 8. Evaluation
##### Method names:
* eval_performance

##### 8. eval_performance - PERSONAL IMPLEMENTATION - Run the following code cell:

```python
# 9. eval_performance (personal implementation info)
# ====>>> YOUR ACTIONS: RUN ONLY

try:    
    student_name, student_id = my_name(), my_id()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)
except Exception as e:
    print ('You probably did not implement student-info functions, \nerror Message:',str(e))
    raise
# ---------------------------
                    
print ("The evaluation metric you need to implement is: '" + assignment_params['eval_metric'] + "'")
```

Output:
```
The evaluation metric you need to implement is: 'recall'
```

<b>8. eval_performance - Permitted and illegal methods</b>:<br />
<div class="alert alert-success">
&#x2714; You can use any function that calculates basic vector or scalar mathematical function
</div>
<div class="alert alert-danger">
&#x2718; DON'T USE any function which performs any of the evaluation methods<br/>
&#x2718; DON'T USE any function which performs a closer step towards the evaluation methods
</div>

### 8. eval_performance
<pre>
input parameters:
- y_predicted -  a series of containing all predicted class values per test instance
- y_test - a series of containing all actual class values per test instance
------------
The function should return a scalar (number) representing the performance evaluation metric (ASSIGNED TO YOU), corresponding to the input parameters. 

<b>Important Notes - <u>for those who need to implement precision or recall</u></b>:
  * <u>The <b>positive</b> category value</u> refers to <u>the value <b>1</b></u>.
  * <u>The <b>negative</b> category value</u> refers to <u>the value <b>0</b></u>.
-----
An example of a return statement:</pre>
<code>    return evaluate_value</code>

##### 8. eval_performance - OPTIONAL ASSISTANCE  - Implement & Run the following code cell (IF NEEDED):

<!-- grade_id: assistance8-eval_performance, points: 0, grade: false, solution: true -->
```python
#8. eval_performance (optional assitance for your solution)
# ====>>> YOUR ACTIONS: (optional) IMPLEMENT & RUN 

# ---- Add assistance code here IF NEEDED:
# YOUR CODE HERE

```

##### 8. eval_performance - Implement and Run the following code cell:

<!-- grade_id: solution8-eval_performance, points: 0, grade: false, solution: true -->
```python
# 8. eval_performance (your solution)
# ====>>> YOUR ACTIONS: IMPLEMENT & RUN 

def eval_performance(y_predicted,y_test):
    length=0
    if(len(y_predicted)==len(y_test)):
        length=len(y_predicted)
    else:
        length=min(len(y_predicted),len(y_test))
    tp=0
    fn=0
    for i in range(0,length):
        if((y_predicted.iloc[i]==1) & (y_test.iloc[i]==1)):
            tp+=1
        if((y_predicted.iloc[i]==0) & (y_test.iloc[i]==1)):
            fn+=1
    evaluate_value=tp/(tp+fn)
    return evaluate_value
    
```

##### 8. eval_performance - OPTIONAL SELF TEST - Implement & Run the following cell (IF NEEDED):

<!-- grade_id: cell-f3c2cfcdcc0fc1da, points: 0, grade: false, solution: true -->
```python
# 8. eval_performance (optional self testing)
# ====>>> YOUR ACTIONS: (optional) IMPLEMENT & RUN 

file_name = 'data' + os.sep + 'Students-Performance_scaled_100.csv'
y_col     = 'category'
student_name, student_id = my_name(), my_id()
assignment_params = get_assignment_params(student_id)
indx = [3,1,6,5,7,44,32,14,10,11]
y_pred = pd.Series([1,1,1,0,1,0,1,0,1,0],index=indx)
y_test = pd.Series([0,0,1,1,1,0,0,0,1,0],index=indx)
eval_res = eval_performance(y_pred, y_test)
# --- add additional code to check your code if needed:
# YOUR CODE HERE


```

##### 8. eval_performance - GRADED TESTS -  Run the following (2) code cells:

<!-- grade_id: test8-1_eval_performance, points: 0.1, grade: true, solution: false -->
```python
# 8. eval_performance (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("8 'eval_performance' - Test 1 (name: test8-1_eval_performance, points: 0.1) - Sanity")

print ("Part 9. - Test 1 (name: test9-1_eval_performance, points: 0.1) - Sanity")
print ("\t--->Testing the implementation of 'eval_performance' ...")

file_name = 'data' + os.sep + 'Students-Performance_numeric.csv'

try:    
    student_name, student_id = my_name(), my_id()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)
    print ("\t\tTesing a '" + assignment_params['eval_metric'] + "' evaluation method ...\n")
    indx = [3,1,6,5,7,44,32,14,10,11]
    y_pred = pd.Series([1,1,1,0,1,0,1,0,1,0],index=indx)
    y_test = pd.Series([0,0,1,1,1,0,0,0,1,0],index=indx)
    eval_res = eval_performance(y_pred, y_test)
    
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

print ("Good Job!\nYou've passed the 1st test for the 'eval_performance' function implementation :-)")
```

Output:
```
8 'eval_performance' - Test 1 (name: test8-1_eval_performance, points: 0.1) - Sanity
Part 9. - Test 1 (name: test9-1_eval_performance, points: 0.1) - Sanity
	--->Testing the implementation of 'eval_performance' ...
		Tesing a 'recall' evaluation method ...

Good Job!
You've passed the 1st test for the 'eval_performance' function implementation :-)
```

<!-- grade_id: test8-2_eval_performance, points: 0.9, grade: true, solution: false -->
```python
# 8. eval_performance (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("8. 'eval_performance' - Test 2 (name: test8-2_eval_performance, points: 0.9)")

file_name = 'data' + os.sep + 'Students-Performance_scaled_100.csv'
y_col     = 'category'

try:    
    student_name, student_id = my_name(), my_id()
    assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"
    assignment_params = get_assignment_params(student_id)
    print ("\t\tTesing a '" + assignment_params['eval_metric'] + "' evaluation method ...\n")
    indx = [3,1,6,5,7,44,32,14,10,11]
    y_pred = pd.Series([1,1,1,0,1,0,1,0,1,0],index=indx)
    y_test = pd.Series([0,0,1,1,1,0,0,0,1,0],index=indx)
    eval_res = eval_performance(y_pred, y_test)
    
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise
    
if assignment_params['eval_metric'] == 'precision':
    assert eval_res == 0.5, 'Wrong precision value'
elif assignment_params['eval_metric'] == 'recall':
    assert eval_res == 0.75, 'Wrong recall value'
elif assignment_params['eval_metric'] == 'accuracy':
    assert eval_res == 0.6, 'Wrong accuracy value'
else: # if assignment_params['eval_metric'] == 'error_rate':
    assert eval_res == 0.4, 'Wrong error-rate value'
    
print ("Good Job!\nYou've passed the 2nd test for the 'eval_performance' function implementation :-)")
```

Output:
```
8. 'eval_performance' - Test 2 (name: test8-2_eval_performance, points: 0.9)
		Tesing a 'recall' evaluation method ...

Good Job!
You've passed the 2nd test for the 'eval_performance' function implementation :-)
```

```python

```